```python
# Initialize Otter
import otter
grader = otter.Notebook("Week10LabPart2.ipynb")
```

# Week 10 Lab Part 2

Please complete this notebook by filling in the cells provided. Before you begin, execute the cell at the top to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

## 1. More Practice with Functions

### **Question 1.1**
[5pts]
Define a function named `coin_flip()` that **returns** "Heads" or "Tails". The function should take NO arguments. Inside the function, generate a random value 1 or 0. Assume the value 1 represents "Heads" and 0 represents "Tails". Call your function to test it.

```python
# your code here
...
```

```python
grader.check("q1_1")
```

### **Question 1.2**

[5 pts] In computer animation, a "jiffy" is commonly defined as 1/100th of a second. Define a function named jiffies_to_seconds that takes the number of "jiffies" as an argument, and **returns** the number of seconds. Call your function to test it.

```python
# your code here
...
```

```python
grader.check("q1_2")
```

### **Question 1.3**

[5 pts] Write a function called `strlist_to_numlist()` that uses a list with *string elements* of numbers as an argument. An example argument would be: ['1', '2', '3']. Your function will convert it to a list with *integer elements*. **Return** the list of integer elements. Call your function to test it and store the result in the variable called `nums`

```python
# your code here
...
```

```python
grader.check("q1_3")
```

### **Question 1.4**
[5pts] Make a function called `tickets()` that **prints** out tickets for each seat in an auditorium. The function should take two arguments: a list of strings representing the row letters, and a list of integers representing the seatnumbers in each row. Call your function to test it.

```python
# your code here
...
```

```python
grader.check("q1_4")
```

## 2. Creating Numpy Arrays

First let's import the numpy package.

```python
import numpy as np
```

Remember, you can create a numpy array using `np.array()` like this:

```python
# create a 1-dimensional array
vector_1D = np.array([1,2,3,4,5])
print(vector_1D)
```

We can check the shape of the array we just made using `np.shape()`

```python
np.shape(vector_1D)
```

Since the shape only gives one value, this is a 1-dimensional (1-D) array. Usually when we are working with vectors and matrices we will want to use 2-D arrays. There are some cases where we can get away with using a 1-D array for vectors. When representing vectors, there is a subtle difference between a 1-D array, and a 2-D array with 1 row. Let's try representing a vector using a 2-D array with 1-row. We will need to use double brackets for this:

```python
#create a 2-D array with only 1 row
row_vector_2D = np.array([[1,2,3,4,5]]) #1 row, 5 columns
print(row_vector_2D)
```

How does this look compared to the 1-D array we created earlier?

```python
print(vector_1D)
```

Again, the 2-D array is surrounded by double brackets, but the 1-D array only has single brackets. Also when we check the shape, the 2-D array has two values and one of the values is 1. The first number represents the number of rows and the second value represents the number of columns. 

```python
np.shape(row_vector_2D)
```

Note how this is different from our 1-D vector:

```python
np.shape(vector_1D)
```

Instead of a horizontal vector, now let's make a vertical vector. We can do this by creating a 2-D array with only 1 column. Note the outer brackets and how each single element row is surrounded by inner brackets and separated by commas:

```python
# create a 2-D array with only 1 column
column_vector_2D = np.array([[1],
                             [2],
                             [3],
                             [4],
                             [5]])
print(column_vector_2D)
```

Now we see that when we check the shape, the first value is the number of rows and the second value is the number of columns.

```python
np.shape(column_vector_2D)
```

What if we want to create a matrix? The outer brackets represent the matrix. The inner brackets represent each row, which are separated by a comma. Here we make a matrix with 3 rows and 2 columns:

```python
# create a 2-d array with 3 rows and 2 columns
matrix_3x2 = np.array([[1,2],
                       [3,4],
                       [5,6]])
print(matrix_3x2)
```

Let's check the shape:

```python
print(np.shape(matrix_3x2))
```

3 rows and 2 columns as expected!

Now it is your turn to make a matrix using numpy!

### **Question 2.1**

[2pts] Create a 2-D array name `matrix_2x3` which has 2 rows and 3 columns. You can use whatever numbers you like inside your array. Check to make sure you created a matrix with the correct dimensions using `np.shape()`

```python
# your code here
...
```

```python
grader.check("q2_1")
```

## 3. Indexing and Slicing a 1-D Array

Indexing and Slicing a 1-D array is similar to indexing a list or a tuple.

```python
print('array=',vector_1D)
print('index=',np.array([0, 1, 2, 3, 4]))
print('The value at the zero index of the array is: ',vector_1D[0])
print('The value at the last index of the array is: ',vector_1D[-1])
print('The slice of the array [1:4] starts at index 1 and goes up to but does not include index 4')
print(vector_1D[1:4])
```

### **Question 3.1**

[2pts] Grab the numbers 3 and 4 from the array `vector_1D` using a slice and store them in a new variable called `slice`.

Hint: To start, ask yourself what index represents the location where `vector_1D` = 3?

```python
print('array=', vector_1D)
print('index=', np.array([0,1,2,3,4]))
# your code here
slice = ...
...
```

Slices that start at the very beginning can be written like this

```python
print('array=', vector_1D)
print('index=', np.array([0,1,2,3,4]))
print('The slice of the array [:3] starts at index 0 and goes up to but does not include index 3')
print(vector_1D[:3])
```

Slices that include the very end can be written like this


```python
print(vector_1D[3:])
```

## 4. Indexing and Slicing a 2-D Array

Recall our 3x2 Matrix

```python
print(matrix_3x2)
```

If we want to grab the number 5 from the matrix using indexing, we do it like this:

```python
print(matrix_3x2[2,0]) #Row 2 (third row), Column 0 (first column)
```

### **Question 4.1**

[3 pts] Grab the number 2 from `matrix_3x2` using indexing and store it in the variable `num2`.

Hint: What row is the number 2 located in? What column?

Hint: What index represents that row? That column?

```python
matrix_3x2 = np.array([[1,2],
                       [3,4],
                       [5,6]])
# your code here
num2 = ...
...
```

If we want to get a slice of a 2-D array we need to provide the rows and columns that we want to grab. Say we have a 3x3 matrix and we want to grab the numbers in the upper right of the matrix:

```python
matrix_3x3 = np.array([[1,2,3],
                       [4,5,6],
                       [7,8,9]])
print(matrix_3x3)
print(' ')
print(matrix_3x3[:2,1:]) #row 0 to row 2 (not inclusive of 2), column 1 to last column
```

### **Question 4.2**

[3pts] Grab a slice that looks like this from `matrix_3x3`:

[[5,6],

[8, 9]]

and store it in the variable `matrix_lower_right`.

```python
matrix_3x3 = np.array([[1,2,3],
                       [4,5,6],
                       [7,8,9]])
# your code here
matrix_lower_right = ...
...
```

## 5. Matrix Functions

When we make functions that deal with matrices, we often need nested loops to loop through every element in the matrix.

For example, the following code is a function that creates an *Identity Matrix* that matches the size of a given square matrix. (Remember that the identity matrix must be a square matrix.) 

```python
def identity(matrix):
    m,n = np.shape(matrix) #find the number of rows and columns in the matrix
    I_matrix = np.zeros([m,n]) #create a matrix of the same size with all zeros
    for i in range(n): #Loop over all the rows
        for j in range(n): #Loop over all the columns
            if i==j: #If the row equals the column, you are on the diagonal
                I_matrix[i, j] = 1 #place a 1 on the diagonal
    return I_matrix

# Test the function
I = identity(matrix_3x3)
print(I)
```

### **Question 5.1**

[10pts] Create a function called `scale_mat()` that scales a matrix by using a nested loop to access each element in the matrix and then multiplying it by the scalar. The function should take the scalar (float) and the matrix (np.array) as arguments IN THAT ORDER. Return the scaled matrix. Call your function to test it.

```python
# Your function here. Remember: You need a nested loop!
# Note: You may check your work by simply multiplying the scalar by the numpy array
# (Isn't it nice that numpy does this for you!)
...
```

### **Question 5.2**
[10pts]
Create a function called `add_mats()` that adds two matrices together by using a nested loop to access each element in the matrices and then adds the elements together and places their sum in a new matrix. The function should take two matrices (both np.array) as arguments. Return the new matrix. Call your function to test it.

```python
# Your function here. Remember: You need a nested loop!
# Note: You may check your work by simply doing np_array1 + np_array2
# (Isn't it nice that numpy does this for you!)
...
```

### **Question 5.3**
[10pts]
Create a function called `mult_mats()` that multiplies two matrices together by using a nested loop. The function should take two matrices (both np.array) as arguments. The first argument should be the left matrix of the multiplication and the second argument should be the right matrix of the multiplication. Return the new matrix that stores the matrix-multiplication result. Call your function to test it. 

If the two matrices are not conformable for multiplication have the function **print** 'Not conformable for matrix multiplication' and **return** nothing at all.

Hint: Check out the General Forumla on Slide 40 from the [Lecture Slides]( https://docs.google.com/presentation/d/1Ll5RWw0MhV33ob-4cmfC4bk1n5ksUm39ZKdA17q4K_0/edit?usp=sharing ).



```python
# your function here
# Note: You may check your work by simply doing np_array1 @ np_array2
# (Isn't it nice that numpy does this for you!)
...
```

## Rotating a vector 

We can use matrix multiplication to rotate a vector. 

If `v` is our vector and `R` is our counterclockwise rotation matrix, then `v` matrix-multiplied by `R` will give us the rotated vector.

We can find `R` using the equation provided in this image, where theta is the angle amount we wish to rotate the vector by

<img src="https://lh3.googleusercontent.com/pw/AP1GczOyapUz6DgVw7S_WqjU61-Bvqd2rWD9h-gBeGEoxrs--xlI9Jb8vf1-6Vg9pzOqTK9gDhOoJrYoT4Bb0QmYjcYU-yEr-DeYNZoIeiItn1j4hQbr8IQMa5Or9ag1da6-wXCg6nrz1m_VIM9YBPuPWehE=w1249-h672-s-no-gm?authuser=0"
     width=800 />

Say we wanted to rotate the column vector [1,0] by 90 degrees:


```python
import math
v = np.array([[1],
              [0]])
theta_degrees = 90 #degrees
theta_radians = math.radians(theta_degrees)
R = np.array([[math.cos(theta_radians), -math.sin(theta_radians)],
              [math.sin(theta_radians),  math.cos(theta_radians)]])
v_rotated = R@v
print(np.round(v_rotated))
np.shape(v)
```

We can plot the original vector and the rotated vector we just calculated

```python
import matplotlib.pyplot as plt
plt.arrow(0, 0, v[0,0], v[1,0], length_includes_head=True, head_width=0.1, head_length=0.2, color='blue',zorder=2,label='v')
plt.arrow(0, 0, v_rotated[0,0], v_rotated[1,0], length_includes_head=True, head_width=0.1, head_length=0.2, color='red',zorder=3,label='v rotated by '+str(theta_degrees)+' degrees')
plt.grid(True,zorder=0) #turn on grid lines, put them underneath everything else (zorder)
plt.legend()
```

### **Question 5.4**
[10pts]
Create a function called `rotate_vec()` that rotates a two element column vector counterclockwise using matrix multiplication.


You may use the @ symbol to perform the matrix multiplication

Remember: vector_rotated = rotation_matrix @ vector

The function should take the two element column vector (2-D numpy array with shape (2,1)) and the rotation angle in degrees (scalar) as arguments, IN THAT ORDER. 
It should return the rotated column vector. Call your function AND then plot the original and rotated vector to test it to make sure your function works as expected.

Hint: The example code above is useful.

```python
# your code here
...
```

You're done with the lab!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [2])
```
>>> coin_flip() == 'Heads' or coin_flip() == 'Tails'
True
```


Tests for q1_2 (points: [2])
```
>>> jiffies_to_seconds(100) == 1.0
True
```


Tests for q1_3 (points: [1, 1])
```
>>> nums = strlist_to_numlist(['1', '2', '3'])
>>> nums == [1, 2, 3]
True
```
```
>>> print(type(nums[0]))
<class 'int'>
```


Tests for q1_4 (points: [5])
```
>>> row = ['A', 'B', 'C']
>>> seat = [1, 2, 3, 4, 5, 6]
>>> tickets(row, seat)
A 1
A 2
A 3
A 4
A 5
A 6
B 1
B 2
B 3
B 4
B 5
B 6
C 1
C 2
C 3
C 4
C 5
C 6
```


Tests for q2_1 (points: [2])
```
>>> np.shape(matrix_2x3) == (2, 3)
True
```


Tests for q3_1 (points: [])


Tests for q4_1 (points: [])


Tests for q4_2 (points: [])


Tests for q5_1 (points: [])


Tests for q5_2 (points: [])


Tests for q5_3 (points: [])


Tests for q5_4 (points: [])
