```python
# Initialize Otter
import otter
grader = otter.Notebook("Week9Lab.ipynb")
```

# Lab Week 9: Numpy, Plotting, Vectors, & Functions!

Please complete this notebook by filling in the cells provided. Before you begin, execute the cell at the top to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

```python
# make sure you run this AND THE CELL AT THE VERY TOP!!!! SCROLL UP!!
# Run this cell to set up the notebook, but please don't change it.
from datascience import *
```

# 1. Functions

### Question 1.1

Define a user function called `welcome()` that **prints** "Welcome to Lab 9". Then call the function.

```python
# your code here
...
```

```python
grader.check("q1_1")
```

### Question 1.2

Define a user function `date_format()` with 3 parameters: `day_num`, `month_word`, and `year_full_num`. The `month_word` parameter should be the full word for the month. The function should then **print** a numeric date with slashes in the format: M/D/Y, where only the last two digits of the year are shown.

Call your function using *positional arguments* to test it.
For example: `date_format(3,'January',2019)` should output 1/3/19. Make sure your function works for any date between the start of year 1000 A.D. and now. (Assume the user knows what century they are referencing.)

```python
# your code here
...
```

```python
grader.check("q1_2")
```

<!-- BEGIN QUESTION -->

### Question 1.3
Add a docstring into the function you defined in Question 1.2 that describes the type and format of each parameter and the format of the function output. Use the help() function to check to make sure your docstring displays how you'd like.

```python
# your code here
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.4

Call your function from Question 1.2 using *keyword arguments*. Assign the `month_word` parameter using the **first argument**. Assign the `day_num` parameter using the **second argument**. Assign the `year_full_num` parameter using the **third argument**.

```python
# your code here
...
```

<!-- END QUESTION -->

### Question 1.5

Define a user function `driving_cost()` with 3 parameters `miles_per_gallon`, `dollars_per_gallon`, and `miles_driven`, that **returns** the dollar cost to drive those miles. Your function should round the dollar cost to the second decimal place using the `round()` function before returning the value. Hint: you can read about the `round()` function using `help(round)`.

Call your function to make sure it works. For example: `c = driving_cost(20.0, 5.1799, 50.0)` should store 12.95 in `c`. Print `c` to check.



```python
# your code here
...
```

```python
grader.check("q1_5")
```

### Question 1.6
Write a user function `rectangle_info()` with parameters `length` and `width`. The function should **return** the rectangle perimeter AND the rectangle area IN THAT ORDER.

Call your function and store each return value in a separate variable by **"unpacking"** the returned tuple.

```python
# your code here
...
```

```python
grader.check("q1_6")
```

# 2. Vectors

```python
# make sure to run this code!
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

## 2.1 Plotting Vectors

```python
u = np.array([2, 4]) #1d array
print(u)

plt.arrow(0, 0, u[0], u[1], length_includes_head=True, head_width=0.2, head_length=0.4, color='blue',zorder=2)
plt.xlim([0, 6]) #x axis limits
plt.ylim([0, 6]) #y axis limits
plt.grid(True,zorder=1) #turn on grid lines, put underneath the arrow using a lower zorder
plt.show()
```

<!-- BEGIN QUESTION -->

**Play around with the code that creates the figure above.**

### Question 2.1.1

The first two arguments of the `plt.arrow()` function are zero. Describe what happens if you change them both to 1.



_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1.2

Describe what happens if you make `length_includes_head=False`?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1.3

Describe what happens if you change the `head_width` and the `head_length` values?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1.4

***Within the `plt.arrow()`*** function, describe what happens if you change `zorder=2` to `zorder=0`?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1.5

Change `plt.xlim([0, 6])` to `plt.xlim([-6, 6])`. Describe what happens to the plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1.6

Change `plt.grid(True,zorder=1)` to `plt.grid(False)`. Describe what happens to the plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2.2 Magnitude of a vector
Suppose we have a vector with n elements:

$\vec{v} = (v_1,v_2,...,v_n)$ in ${\rm I\!R}^n$

The magnitude of the vector can be found as,

$||\vec{v}|| = \sqrt{v_1^2 + v_2^2 + ...v_n^2}$

### Question 2.2.1

Create a USER FUNCTION called `mag` that calculates the magnitude of a vector. The function should take *one argument that is either a list or an array of numbers* which represents a vector. It should **return** one value which is the magnitude. Use a LOOP to calculate the summation under the root.

Make sure your function correctly returns the magnitude for any size vector.

```python
#Your function here
...
```

```python
#test your function on this 2D vector
u = np.array([3,4])
print(mag(u))
#is the resulting magnitude what you expect?
```

```python
#test your function on this random 20 element vector
u = np.random.randint(0,100,20)
print("u=",u)
print(mag(u))
#does it work even for a large number of elements like this?
```

```python
grader.check("q2_2_1")
```

## 2.3 Direction of a vector

<img src="https://lh3.googleusercontent.com/pw/ADCreHdSG3G7W2bbYr7SdhANlnvaU_Vvjv6Hb-63UbWVXYDASIAmQbCc1Xt1L0Y-wccyxz5NVWWXtTClxSeosV5Ifzk1wJT88MjE9PWIzzBiMTa2KheqC4_dFxaxytn8S1LE5wK5gvjuDqKNWugvlR7BT-Wf=w484-h463-s-no-gm?authuser=0" width="350">

Suppose we have a 2-D vector :

$\vec{v} = (v_1,v_2)$

Using geometry,

$tan(\theta) = \frac{v_2}{v_1}$

So

$\theta = tan^{-1}(\frac{v_2}{v_1})$

Note that $\theta$ is defined as the counterclockwise angle from the x-axis, such that a vector in the bottom-right quadrant (4th quadrant) has $\theta$ as depicted below:

<img src="https://lh3.googleusercontent.com/pw/ADCreHe37CmfFVVRHFSIYPkdQ2I_g-QuW2cLxwWrZ0xZUzEE032F_96LoSFeutZVNKgkfnZFZQyadyn37tLoISCFoOJ_GtAa4T9zv90QOAfVpBOcUh2zd-5yKO0VV--9G3c886bSkC3JE8fkdwIuWncwWqOF=w538-h520-s-no-gm?authuser=0" width="350">




<!-- BEGIN QUESTION -->

### Question 2.3.1

Create four 2D vectors, where each vector is in a different quadrant:

<img src="https://lh3.googleusercontent.com/pw/ADCreHemWnJcdihmnkpSHHvlO9ouQEICiuNgPQqycTB-CuJwzDjxyL0LoMWlVzI3kU65m3Oextt-mn4MpAsSw5UUUA3Cdcxl6TQqEn1ShkEY4kBRfYrV3q8VJ4TdK3WUXQIrXIQ2aB4PUY-CGHlunjybaoKY=w290-h258-s-no-gm?authuser=0" width="350">

  

```python
#first quadrant
u1 = np.array([2,2])
#second quadrant
u2 = np.array([-2,2])
#third quadrant
#your code here
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.3.2

Calculate the direction of each of the vectors you made using the `np.arctan()` inverse tangent function, and convert to degrees using the `np.degrees()` function. Add a "fudge factor" when needed to get the angle to match the geometric definition of vector direction in the schematics above.  

```python
#first quadrant
print('first quadrant')
#first try using np.atan()
theta1 = np.arctan(u1[1]/u1[0])
print(np.degrees(theta1))
print('')

#second quadrant
print('second quadrant')
#first try using np.atan()
theta2 = np.arctan(u2[1]/u2[0])
# Hrm this isn't what I expect
print(np.degrees(theta2))
# Add a fudge factor to get it to be what I think it should be
fudge = 180
print(np.degrees(theta2)+fudge)
print('')

#your code here
...
```

<!-- END QUESTION -->

### Question 2.3.3
Write a USER FUNCTION called `vecdir` using IF-ELIF-ELSE branching that **returns** the direction (theta) of a 2D vector using the code you wrote above. Your function should take a list or a numpy array as an argument.

Test your user function using the vectors you created above, AND using vectors that lie on the boundaries between each quadrant. 

```python
#your user function code here
...
```

```python
grader.check("q2_3_3")
```

## 2.4 Vector Addition

Numpy arrays make vector addition easy:

```python
u = np.array([2, 4]) #1d array
print(u)
v = np.array([5, 3]) #1d array
print(v)

print("u+v=",u+v)
```

<!-- BEGIN QUESTION -->

### Question 2.4.1

Using the vectors defined in the code above, create a plot showing the addition of vectors $\vec{u}$ and $\vec{v}$, and the resultant vector $\overrightarrow{u+v}$. Use the vectors from the code above.

Place the tail of $\vec{u}$ at the origin. Place the tail of $\vec{v}$ at the head of $\vec{u}$. Place the tail of $\overrightarrow{u+v}$ at the origin.  

Hint: Refer to your answer in Question 2.1.1. Make each vector a different color and show all your vectors on the same plot.

```python
#Your plotting code here
...
```

<!-- END QUESTION -->

## 2.5 Scaling a vector

Numpy arrays make scaling a vector easy too:

```python
u = np.array([2, 4]) #1d array
print(u)
c = 2

print("2u=",c*u)
```

<!-- BEGIN QUESTION -->

### Question 2.5.1

Using the vectors defined in the code above, create a plot showing the original vector $\vec{u}$ and scaled vector $2\vec{u}$. Place the tails of both vectors at the origin. Hint: Refer to your answer in Question 2.1.4. Make each vector a different color and show all your vectors on the same plot.

```python
#your plotting code here

...
```

<!-- END QUESTION -->

## 2.6 Unit vectors

A unit vector has a magnitude of 1.

To find a unit vector in the same direction as $\vec{v}$:

$\frac{1}{||\vec{v}||}\vec{v}$

<!-- BEGIN QUESTION -->

### Question 2.6.1
Using your magnitude user function from Question 2.2.1, find a unit vector in the direction of $\vec{v}$. Plot the unit vector on top of the original vector. Place both of the vector tails at the origin.

```python
v = np.array([2, 4]) #1d array
# plot this vector and a unit vector in the same direction
...
```

<!-- END QUESTION -->

You're done with the lab!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [2])
```
>>> welcome()
Welcome to Lab 9
```


Tests for q1_2 (points: [2])
```
>>> date_format(3, 'January', 2019)
1/3/19
```


Tests for q1_5 (points: [2])
```
>>> driving_cost(20.0, 5.1799, 50.0) == 12.95
True
```


Tests for q1_6 (points: [2])
```
>>> rectangle_info(18, 8) == (52, 144)
True
```


Tests for q2_2_1 (points: [5])
```
>>> mag([3, 4]) == 5
True
```


Tests for q2_3_3 (points: [2, 2])
```
>>> u1 = np.array([2, 2])
>>> vecdir(u1) == 45.0
True
```
```
>>> u5 = np.array([3, 0])
>>> vecdir(u5) == 0.0 or vecdir(u5) == 360.0
True
```
